# <span style="color:seagreen;">Bi 410/510 Project 9:  Pandas</span>

##### Due 11:00 P.M. March 20, 2017

##### Reading:

* lecture notes from Mar 13
* _CSB_ Chapter 6 (skim 6.1--6.2, read 6.3)
* documentation online at http://pandas.pydata.org (esp. "10 minutes to Pandas")

###  <span style="color:seagreen">Instructions</span>

The final programming project is a quick introduction to Pandas, a Python library widely used for data analysis. 

The data set to use for this project is in a file named `clusters.csv` that you can download from Canvas.

The exercises in this notebook are less structured than previous programming exercises.  One of the goals this week is to help you explore the Pandas documentation to see what's possible and then to take what you learn by reading the documentation and apply it to our data set.

###  <span style="color:seagreen">Import the Library</span>

Execute the following code cell each time you open this notebook.

```python
import pandas as pd
import numpy as np
```

If you get a "no module" error you need to install one or both of the libraries -- see the instructions in the Mar 13 lecture notes.

### <span style="color:seagreen">1. &nbsp; Read the Data </span>

Use the library function named `read_csv` to load the data from `clusters.csv` into a new data frame named `clusters` and then display the contents of the frame.  You can use a `print` statement, or, if you ask Python to display the table, it will show up looking something like this in your notebook:

![frame](http://pages.uoregon.edu/conery/Bi410/Pandas/frame.png)


<!-- grade_id: clusters, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
pass
```

### <span style="color:seagreen">2. &nbsp; Column Names </span>

The three data columns are the expected number of sequences for each species, the number of clusters identified by `vsearch`, and the number identified by `swarm`.

Write a Python expression that prints the names of the columns in the data frame.

<!-- grade_id: columns, points: 10, grade: true, solution: true -->
```python
# YOUR CODE HERE
pass
```

### <span style="color:seagreen">3. &nbsp; Column Statistics </span>

Write a for loop that prints the number of clusters and the total number of sequences in each column.  Your output should look something like this:

```
expected   21  100149
vsearch    14   95939
swarm      23   95953
```

(It's OK if you don't have the exact same format, or if some of your values are floats).

<!-- grade_id: stats, points: 20, grade: true, solution: true -->
```python
# YOUR CODE HERE
pass
```

### <span style="color:seagreen">4. &nbsp; Differences </span>

Write a function that compares two columns and returns the sum of the differences between them.

Note that columns might not have the same number of values since some columns have NaN ("not a number") entries.  For example, consider a frame named `df` with these two columns:
```
       A     B
      ===   ===
       24    20
       17    16
        8     4
        3   NaN
        1   NaN
```

The simplest way to define the number of differences is to use only the rows where both columns have values, in which case the count is 4 + 1 + 4 = 9.

Your function should be named `diffs`, and it should take two columns as arguments.  Each column in a frame is a "series" object, so the straightforward approach to computing the differences would be to make a new series by subtracting one column from the other and then computing the sum of the result.

**Hint:** &nbsp; If `a` and `b` are series objects, `a - b` is a new series made by subtracting `b` from `a`.  Can you figure out how to find the sum of the absolute values of the differences?

<!-- grade_id: similarity, points: 25, grade: true, solution: true -->
```python
def similarity(col1, col2):
    # YOUR CODE HERE
    pass
```

Our "autograder" will test your function by asking it to compute the differences between the expected counts and those found by `vsearch` and `swarm`.

<!-- grade_id: vsearch_test, points: 5, grade: true, solution: false -->
```python
assert similarity(clusters['expected'], clusters['vsearch']) == 4103
```

<!-- grade_id: swarm_test, points: 5, grade: true, solution: false -->
```python
assert similarity(clusters['expected'], clusters['swarm']) == 5432
```

### <span style="color:seagreen">5. &nbsp; Visualization </span>

Use the code cell below to type a command that plots the distribution of cluster sizes.  Feel free to experiment with different visualzation techniques defined in Pandas.

Here are two examples.  The first is a "bar plot" that shows all the clusters:

![all](http://pages.uoregon.edu/conery/Bi410/Pandas/plot_all.png)

It's hard to see the differences because the count for the first cluster is so much bigger.  This picture leaves out the first 5 rows of data so we can "zoom in" on the remaining rows:

![all](http://pages.uoregon.edu/conery/Bi410/Pandas/plot_rest.png)

**Note:** &nbsp; To display a visualization in the notebook execute the command in this code cell before you call the plot function:

```python
% matplotlib inline
```

<!-- grade_id: plots, points: 25, grade: true, solution: true -->
```python
# YOUR CODE HERE
pass
```

### <span style="color:seagreen">Extra Credit</span>

If you have some experience with statistics explore different ways to compare columns in the cluster data.  Pandas and Numpy have extensive stats libraries.  Can you figure out how to compute correlation coefficients?  Or do a chi-square test, or some other test for how similar the columns are?

```python

```